```python
# Initialize Otter
import otter
grader = otter.Notebook("hw02.ipynb")
```

# Homework 2: Collecting Data; Arrays and Tables

Please complete this notebook by filling in the cells provided.
 
Throughout this homework and all future ones, please be sure to not re-assign variables throughout the notebook! For example, if you use `max_temperature` in your answer to one question, do not reassign it later on. 

<!-- BEGIN QUESTION -->

What is your name?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

## Part 1: Collecting Data

Data could be collected in various ways:
* a survey
* automatic collection of web traffic
* an experiment
* many other ways

When collected, it could be stored in various file formats.
Sometimes the file format is especially good for that type of data, and other times it is a matter of convenience.

<!-- BEGIN QUESTION -->

#### Question 1

Peruse the [the State of Ohio's Data Portal](https://data.ohio.gov) and locate a data file in CSV format.
* Give the URL of the description of this data file. 
* Based on the description, explain (or speculate on) how this data was collected.
* Download and open the file. Copy and paste the first 3 lines of the file here.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

#### Question 2

Peruse the [the State of Ohio's Data Portal](https://data.ohio.gov) and locate a **data** file **not** in CSV format.
(Make sure that it is a data file, rather than a presentation.)

* Give the URL of the description of this data file. 
* Based on the description, explain (or speculate on) how this data was collected.
* Download the file. Either:
  * convert it to CSV and copy and paste the first 3 lines of the file here, or
  * explain why it cannot (or should not) be converted to CSV.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

## Part 2: Arrays and Tables

(This part of the homework is based on [Data 8 hw 02](https://github.com/data-8/materials-sp20/tree/master/materials/sp20/hw/hw02).)

```python
# Don't change this cell; just run it.
# When you log-in please hit return (not shift + return) after typing in your email
import numpy as np
from datascience import *

```

### 2.1. Creating Arrays


**Question 1.** Make an array called `weird_numbers` containing the following numbers (in the given order):

1. -2
2. the sine of 1.2
3. 3
4. 5 to the power of the cosine of 1.2

*Hint:* `sin` and `cos` are functions in the `math` module.

*Note:* Python lists are different/behave differently than numpy arrays. In Data 8, we use numpy arrays, so please make an **array**, not a python list.


```python
# Our solution involved one extra line of code before creating
# weird_numbers.
...
weird_numbers = ...
weird_numbers
```

```python
grader.check("q2_1_1")
```

**Question 2.** Make an array called `book_title_words` containing the following three strings: "Eats", "Shoots", and "and Leaves".


```python
book_title_words = ...
book_title_words
```

```python
grader.check("q2_1_2")
```

Strings have a method called `join`.  `join` takes one argument, an array of strings.  It returns a single string.  Specifically, the value of `a_string.join(an_array)` is a single string that's the [concatenation](https://en.wikipedia.org/wiki/Concatenation) ("putting together") of all the strings in `an_array`, **except** `a_string` is inserted in between each string.

**Question 3.** Use the array `book_title_words` and the method `join` to make two strings:

1. "Eats, Shoots, and Leaves" (call this one `with_commas`)
2. "Eats Shoots and Leaves" (call this one `without_commas`)

*Hint:* If you're not sure what `join` does, first try just calling, for example, `"foo".join(book_title_words)` .


```python
with_commas = ...
without_commas = ...

# These lines are provided just to print out your answers.
print('with_commas:', with_commas)
print('without_commas:', without_commas)
```

```python
grader.check("q2_1_3")
```

### 2.2. Indexing Arrays


These exercises give you practice accessing individual elements of arrays.  In Python (and in many programming languages), elements are accessed by *index*, so the first element is the element at index 0.

*Note:* Please don't use bracket notation when indexing (i.e. `arr[0]`), as this can yield different data type outputs than what we will be expecting. This can cause you to fail an autograder test.

**Question 1.** The cell below creates an array of some numbers.  Set `third_element` to the third element of `some_numbers`.


```python
some_numbers = make_array(-1, -3, -6, -10, -15)

third_element = ...
third_element
```

```python
grader.check("q2_2_1")
```

**Question 2.** The next cell creates a table that displays some information about the elements of `some_numbers` and their order.  Run the cell to see the partially-completed table, then fill in the missing information (the cells that say "Ellipsis") by assigning `blank_a`, `blank_b`, `blank_c`, and `blank_d` to the correct elements in the table.


```python
blank_a = ...
blank_b = ...
blank_c = ...
blank_d = ...
elements_of_some_numbers = Table().with_columns(
    "English name for position", make_array("first", "second", blank_a, blank_b, "fifth"),
    "Index",                     make_array(blank_c, 1, 2, blank_d, 4),
    "Element",                   some_numbers)
elements_of_some_numbers
```

```python
grader.check("q2_2_2")
```

**Question 3.** You'll sometimes want to find the *last* element of an array.  Suppose an array has 142 elements.  What is the index of its last element?


```python
index_of_last_element = ...
```

More often, you don't know the number of elements in an array, its *length*.  (For example, it might be a large dataset you found on the Internet.)  The function `len` takes a single argument, an array, and returns the `len`gth of that array (an integer).

**Question 4.** The cell below loads an array called `president_birth_years`.  Calling `.column(...)` on a table returns an array of the column specified, in this case the `Birth Year` column of the `president_births` table. The last element in that array is the most recent birth year of any deceased president. Assign that year to `most_recent_birth_year`.


```python
president_birth_years = Table.read_table("president_births.csv").column('Birth Year')

most_recent_birth_year = ...
most_recent_birth_year
```

**Question 5.** Finally, assign `sum_of_birth_years` to the sum of the first, tenth, and last birth year in `president_birth_years`.


```python
sum_of_birth_years = ...
```

### 2.3. Basic Array Arithmetic


**Question 1.** Multiply the numbers 42, 4224, 42422424, and -250 by 157. Assign each variable below such that `first_product` is assigned to the result of $42 * 157$, `second_product` is assigned to the result of $4224 * 157$, and so on. 

For this question, **don't** use arrays.


```python
first_product = ...
second_product = ...
third_product = ...
fourth_product = ...
print(first_product, second_product, third_product, fourth_product)
```

**Question 2.** Now, do the same calculation, but using an array called `numbers` and only a single multiplication (`*`) operator.  Store the 4 results in an array named `products`.


```python
numbers = ...
products = numbers * 157
products
```

**Question 3.** Oops, we made a typo!  Instead of 157, we wanted to multiply each number by 1577.  Compute the correct products in the cell below using array arithmetic.  Notice that your job is really easy if you previously defined an array containing the 4 numbers.


```python
correct_products = ...
correct_products
```

**Question 4.** We've loaded an array of temperatures in the next cell.  Each number is the highest temperature observed on a day at a climate observation station, mostly from the US.  Since they're from the US government agency [NOAA](noaa.gov), all the temperatures are in Fahrenheit.  Convert them all to Celsius by first subtracting 32 from them, then multiplying the results by $\frac{5}{9}$. Make sure to **ROUND** the final result after converting to Celsius to the nearest integer using the `np.round` function.


```python
max_temperatures = Table.read_table("temperatures.csv").column("Daily Max Temperature")

celsius_max_temperatures = ...
celsius_max_temperatures
```

```python
grader.check("q2_3_4")
```

**Question 5.** The cell below loads all the *lowest* temperatures from each day (in Fahrenheit).  Compute the size of the daily temperature range for each day.  That is, compute the difference between each daily maximum temperature and the corresponding daily minimum temperature.  **Pay attention to the units, give your answer in Celsius!** Make sure **NOT** to round your answer for this question!


```python
min_temperatures = Table.read_table("temperatures.csv").column("Daily Min Temperature")

celsius_temperature_ranges = ...
celsius_temperature_ranges
```

### 2.4. World Population


Remember that the tests from this point on will **not** necessarily tell you whether or not your answers are correct.

The cell below loads a table of estimates of the world population for different years, starting in 1950. The estimates come from the [US Census Bureau website](https://www.census.gov/en.html).

```python
world = Table.read_table("world_population.csv").select('Year', 'Population')
world.show(4)
```

The name `population` is assigned to an array of population estimates.

```python
population = world.column(1)
population
```

In this question, you will apply some built-in Numpy functions to this array. Numpy is a module that is often used in Data Science!

<img src="array_diff.png" style="width: 600px;"/>

The difference function `np.diff` subtracts each element in an array from the element after it within the array. As a result, the length of the array `np.diff` returns will always be one less than the length of the input array.

<img src="array_cumsum.png" style="width: 700px;"/>

The cumulative sum function `np.cumsum` outputs an array of partial sums. For example, the third element in the output array corresponds to the sum of the first, second, and third elements.

**Question 1.** Very often in data science, we are interested understanding how values change with time. Use `np.diff` and `np.max` (or just `max`) to calculate the largest annual change in population between any two consecutive years.


```python
largest_population_change = ...
largest_population_change
```

**Question 2.** What do the values in the resulting array represent (choose one)?

```python
np.cumsum(np.diff(population))
```

1) The total population change between consecutive years, starting at 1951.

2) The total population change between 1950 and each later year, starting at 1951.

3) The total population change between 1950 and each later year, starting inclusively at 1950.


```python
# Assign cumulative_sum_answer to 1, 2, or 3
cumulative_sum_answer = ...
```

```python
grader.check("q2_4_2")
```

### 2.5. Old Faithful


Old Faithful is a geyser in Yellowstone that erupts every 44 to 125 minutes (according to [Wikipedia](https://en.wikipedia.org/wiki/Old_Faithful)). People are [often told that the geyser erupts every hour](http://yellowstone.net/geysers/old-faithful/), but in fact the waiting time between eruptions is more variable. Let's take a look.

**Question 1.** The first line below assigns `waiting_times` to an array of 272 consecutive waiting times between eruptions, taken from a classic 1938 dataset. Assign the names `shortest`, `longest`, and `average` so that the `print` statement is correct.


```python
waiting_times = Table.read_table('old_faithful.csv').column('waiting')

shortest = ...
longest = ...
average = ...

print("Old Faithful erupts every", shortest, "to", longest, "minutes and every", average, "minutes on average.")
```

```python
grader.check("q2_5_1")
```

**Question 2.** Assign `biggest_decrease` to the biggest decrease in waiting time between two consecutive eruptions. For example, the third eruption occurred after 74 minutes and the fourth after 62 minutes, so the decrease in waiting time was 74 - 62 = 12 minutes. 

*Hint 1*: You'll need an array arithmetic function [mentioned in the Data 8 textbook](https://inferentialthinking.com/chapters/05/1/Arrays.html#functions-on-arrays). You have also seen this function earlier in the homework!

*Hint 2*: We want to return the absolute value of the biggest decrease.


```python
biggest_decrease = ...
biggest_decrease
```

```python
grader.check("q2_5_2")
```

**Question 3.** If you expected Old Faithful to erupt every hour, you would expect to wait a total of `60 * k` minutes to see `k` eruptions. Set `difference_from_expected` to an array with 272 elements, where the element at index `i` is the absolute difference between the expected and actual total amount of waiting time to see the first `i+1` eruptions. 

*Hint*: You'll need to compare a cumulative sum to a range. You'll go through `np.arange` more thoroughly in a Lab, but you can read about it in this [textbook section](https://www.inferentialthinking.com/chapters/05/2/Ranges.html).

For example, since the first three waiting times are 79, 54, and 74, the total waiting time for 3 eruptions is 79 + 54 + 74 = 207. The expected waiting time for 3 eruptions is 60 * 3 = 180. Therefore, `difference_from_expected.item(2)` should be $|207 - 180| = 27$.


```python
difference_from_expected = ...
difference_from_expected
```

```python
grader.check("q2_5_3")
```

**Question 4.** Let’s imagine your guess for the next wait time was always just the length of the previous waiting time. If you always guessed the previous waiting time, how big would your error in guessing the waiting times be, on average?

For example, since the first three waiting times are 79, 54, and 74, the average difference between your guess and the actual time for just the second and third eruption would be $\frac{|79-54|+ |54-74|}{2} = 22.5$.


```python
average_error = ...
average_error
```

```python
grader.check("q2_5_4")
```

### 2.6. Tables


**Question 1.** Suppose you have 4 apples, 3 oranges, and 3 pineapples.  (Perhaps you're using Python to solve a high school Algebra problem.)  Create a table that contains this information.  It should have two columns: `fruit name` and `count`.  Assign the new table to the variable `fruits`.

**Note:** Use lower-case and singular words for the name of each fruit, like `"apple"`.


```python
# Our solution uses 1 statement split over 3 lines.
fruits = ...
        ...
        ...
fruits
```

```python
grader.check("q2_6_1")
```

**Question 2.** The file `inventory.csv` contains information about the inventory at a fruit stand.  Each row represents the contents of one box of fruit. Load it as a table named `inventory` using the `Table.read_table()` function. `Table.read_table(...)` takes one argument (data file name in string format) and returns a table.


```python
inventory = ...
inventory
```

```python
grader.check("q2_6_2")
```

**Question 3.** Does each box at the fruit stand contain a different fruit? Set `all_different` to `True` if each box contains a different fruit or to `False` if multiple boxes contain the same fruit.

*Hint:* You don't have to write code to calculate the True/False value for `all_different`. Just look at the `inventory` table and assign `all_different` to either `True` or `False` according to what you can see from the table in answering the question.


```python
all_different = ...
all_different
```

```python
grader.check("q2_6_3")
```

**Question 4.** The file `sales.csv` contains the number of fruit sold from each box last Saturday.  It has an extra column called "price per fruit (\$)" that's the price *per item of fruit* for fruit in that box.  The rows are in the same order as the `inventory` table.  Load these data into a table called `sales`.


```python
sales = ...
sales
```

```python
grader.check("q2_6_4")
```

**Question 5.** How many fruits did the store sell in total on that day?


```python
total_fruits_sold = ...
total_fruits_sold
```

```python
grader.check("q2_6_5")
```

**Question 6.** What was the store's total revenue (the total price of all fruits sold) on that day?

*Hint:* If you're stuck, think first about how you would compute the total revenue from just the grape sales.


```python
total_revenue = ...
total_revenue
```

```python
grader.check("q2_6_6")
```

**Question 7.** Make a new table called `remaining_inventory`.  It should have the same rows and columns as `inventory`, except that the amount of fruit sold from each box should be subtracted from that box's count, so that the "count" is the amount of fruit remaining after Saturday.


```python
remaining_inventory = ...
    ...
    ...
    ...
remaining_inventory
```

```python
grader.check("q2_6_7")
```

## Acknowledgements

<!-- BEGIN QUESTION -->

**Acknowledgements:**
If you received help from any person other than your instructor, then describe here what help you received and from whom. 
If not, then say 'None'.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

It is forbidden to: 
* Use the work or programs from students who took this class (in any of its versions at any university).
* Use websites that claim to have solutions for this class (in any of its versions at any university).
* Directly copy from a classmate.

**Did you do anything forbidden while working on this assignment?**

Assign `did_forbidden` to either `True` or `False`.

```python
did_forbidden = ...
```

```python
grader.check("q_forbidden")
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

After exporting, submit the .zip file in Canvas

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Tests for q2_1_1 (points: None)
```
>>> np.allclose(weird_numbers, np.array([-2.0, 0.93203909, 3.0, 1.79174913]), rtol=0.001, atol=0.001) == True
True
```


Tests for q2_1_2 (points: None)
```
>>> import numpy as np
>>> type(book_title_words) == np.ndarray
True
```
```
>>> not any([',' in text for text in book_title_words])
True
```
```
>>> 'and ' in book_title_words.item(2)
True
```
```
>>> len(book_title_words) == 3
True
```
```
>>> book_title_words.item(0) == 'Eats'
True
```
```
>>> book_title_words.item(1) == 'Shoots'
True
```
```
>>> book_title_words.item(2) == 'and Leaves'
True
```


Tests for q2_1_3 (points: None)
```
>>> with_commas == 'Eats, Shoots, and Leaves'
True
```
```
>>> without_commas == 'Eats Shoots and Leaves'
True
```


Tests for q2_2_1 (points: None)
```
>>> third_element != -10
True
```


Tests for q2_2_2 (points: None)
```
>>> elements_of_some_numbers.column(0).item(2) == 'third'
True
```
```
>>> elements_of_some_numbers.column(0).item(3) == 'fourth'
True
```
```
>>> elements_of_some_numbers.column(1).item(0) == 0
True
```
```
>>> elements_of_some_numbers.column(1).item(3) == 3
True
```


Tests for q2_2_3 (points: None)


Tests for q2_2_4 (points: None)


Tests for q2_2_5 (points: None)


Tests for q2_3_1 (points: None)


Tests for q2_3_2 (points: None)


Tests for q2_3_3 (points: None)


Tests for q2_3_4 (points: None)
```
>>> sum(celsius_max_temperatures) != 356705.0
True
```
```
>>> sum(celsius_max_temperatures) == 1280677.0
True
```
```
>>> len(celsius_max_temperatures) == 65000
True
```
```
>>> celsius_max_temperatures.item(2003) == 20.0
True
```


Tests for q2_3_5 (points: None)


Tests for q2_4_1 (points: None)


Tests for q2_4_2 (points: None)
```
>>> type(cumulative_sum_answer) == int
True
```
```
>>> 1 <= cumulative_sum_answer <= 3
True
```


Tests for q2_5_1 (points: None)
```
>>> shortest <= average <= longest
True
```


Tests for q2_5_2 (points: None)
```
>>> biggest_decrease != 47
True
```
```
>>> 30 <= biggest_decrease < 47
True
```


Tests for q2_5_3 (points: None)
```
>>> difference_from_expected.size == 272
True
```
```
>>> difference_from_expected.item(271) == abs(60 * 272 - sum(waiting_times))
True
```


Tests for q2_5_4 (points: None)
```
>>> 15 <= average_error <= 25
True
```


Tests for q2_6_1 (points: None)
```
>>> fruits.sort(0).column(0).item(0) == 'apple'
True
```
```
>>> fruits.sort(0).column(1).item(0) == 4
True
```


Tests for q2_6_2 (points: None)
```
>>> inventory.sort(0).column(0).item(0) == 25274
True
```


Tests for q2_6_3 (points: None)
```
>>> all_different in {True, False}
True
```


Tests for q2_6_4 (points: None)
```
>>> sales.sort(0).column(0).item(1) == 26187
True
```


Tests for q2_6_5 (points: None)
```
>>> total_fruits_sold > 10
True
```


Tests for q2_6_6 (points: None)
```
>>> 50 <= total_revenue <= 150
True
```


Tests for q2_6_7 (points: None)
```
>>> remaining_inventory.num_columns == 3
True
```
```
>>> remaining_inventory.column('count').item(0) != 45
True
```
```
>>> remaining_inventory.where(1, 'grape').column(2).item(0) == 162
True
```


Tests for q_forbidden (points: None)
```
>>> did_forbidden in [True, False]
True
```
