```python
# Initialize Otter
import otter
grader = otter.Notebook("hw01.ipynb")
```

# Homework 1: Types and Sources of Data and Python Expressions

Please complete this notebook by filling in the cells provided.

<!-- BEGIN QUESTION -->

What is your name?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

## Part 1. Types and Sources of Data

### 1.1 Types of Data

Data feature comes in many types.
Major ones are:
* Nominal (having names): General labels (e.g "Purple Coneflower", "Northern Cardinal", ...)
* Ordinal (having an order): Labels with a natural order (e.g. "Small", "Medium", "Large")
* Quantitative (measuring an amount): Something on a scale (e.g. "1.3kg" or "67%")

Other terms are also used:
* Categorical (having a category/label) includes both Nominal and Ordinal.
* Quantitative can be split into *discrete* (only certain values allowed) and *continuous* (all values in an interval are allowed).

Sometimes the type to call a particular feature is not clear. 
For example, an ordinal feature might be give a number value, which makes it seem like a quantitative feature.

#### Questions

For each of the following features, 
* name what type(s) it is
* give a short justification for your choice.

<!-- BEGIN QUESTION -->

The country of birth of a person.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

The shoe size of a person, using the [United States customary system](https://en.wikipedia.org/wiki/Shoe_size#Customary).

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

The number of hairs on a person's head.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

### 1.2 Sources of Data

A *source* of data could be a person, organization, company, government, etc. 
Some sources provide high-quality, reliable data, while others might provide low-quality, or even intentionally false data.
An example data source, which you may not use when answering questions on this homework, is the [Scalia Lab Laboratory at Ohio University](https://scalialab.com/current/archives/).

<!-- BEGIN QUESTION -->

Identify a source of data on the internet that seems reliable. Give:
* Its URL
* A 2-3 sentence description of what the source is.
* A 2-3 sentence explanation of why you think it is reliable.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

Identify a source of data on the internet that seems **unreliable**. Give:
* Its URL
* A 2-3 sentence description of what the source is.
* A 2-3 sentence explanation of why you think it is unreliable.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

### 1.3 Sets of Data

A single source of data may provide more than one *set* of data.
For example, the [Scalia Lab Laboratory](https://scalialab.com/current/archives/) provides
* [Daily Extremes - Monthly](https://scalialab.com/current/archives/daily-extremes/)
* [Hourly Observations - West State Street](https://scalialab.com/current/archives/hourly/)
* [Hourly Observations - All](https://scalialab.com/current/archives/hourly-all/)

(Looking deeper, these data sets are broken down into smaller data sets, by when they were collected.)	

<!-- BEGIN QUESTION -->

From the reliable data source you identified above, select 2 data sets. 
For each, give:
* Its URL
* A 2-3 sentence description of what the data set is.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

## Part 2. Python Expressions

(This part of the homework is based on [Data 8 hw 01](https://github.com/data-8/materials-sp20/tree/master/materials/sp20/hw/hw01).)

### 2.1. Names and Assignment Statements


**Question 1.** When you run the following cell, Python produces a cryptic error message.

```python
4 = 2 + 2
```

Choose the best explanation of what's wrong with the code, and then assign 1, 2, 3, or 4 to `names_q1` below to indicate your answer.

1. Python is smart and already knows `4 = 2 + 2`.

2. `4` is already a defined number, and it doesn't make sense to make a number be a name for something else. In Python, "`x = 2 + 2`" means "assign `x` as the name for the value of `2 + 2`."

3. It should be `2 + 2 = 4`.

4. I don't get an error message. This is a trick question.


```python
names_q1 = ...
```

```python
grader.check("q2_1_1")
```

**Question 2.** When you run the following cell, Python will produce another cryptic error message.

```python
two = 3
six = two plus two
```

Choose the best explanation of what's wrong with the code and assign 1, 2, 3, or 4 to `names_q2` below to indicate your answer.

1. The `plus` operation only applies to numbers, not the word "two".

2. The name "two" cannot be assigned to the number 3.

3. Two plus two is four, not six.

4. Python cannot interpret the name `two` followed directly by a name that has not been defined.


```python
names_q2 = ...
```

```python
grader.check("q2_1_2")
```

**Question 3.** When you run the following cell, Python will, yet again, produce another cryptic error message.

```python
x = print(5)
y = x + 2
```

Choose the best explanation of what's wrong with the code and assign 1, 2, or 3 to `names_q3` below to indicate your answer.

1. Python doesn't want `y` to be assigned.

2. The `print` operation is meant for displaying values to the programmer, not for assigning values!

3. Python can’t do addition between one name and one number. It has to be 2 numbers or 2 predefined names.


```python
names_q3 = ...
```

```python
grader.check("q2_1_3")
```

### 2.2. Differences between Majors


Berkeley’s Office of Planning and Analysis provides data on numerous aspects of the campus. Adapted from the OPA website, the table below displays the numbers of degree recipients in three majors in the academic years 2008-2009 and 2017-2018.

| Major                              | 2008-2009    | 2017-2018   |
|------------------------------------|--------------|-------------|
| Gender and Women's Studies         |      17      |    28       |
| Linguistics                        |      49      |    67       |
| Rhetoric                           |      113     |    56       |



<!-- BEGIN QUESTION -->

**Question 1.** Suppose you want to find the **biggest** absolute difference between the numbers of degree recipients in the two years, among the three majors.

In the cell below, compute this value and call it `biggest_change`. Use a single expression (a single line of code) to compute the answer. Let Python perform all the arithmetic (like subtracting 49 from 67) rather than simplifying the expression yourself. The built-in `abs` function takes a numerical input and returns the absolute value. The built-in `max` function can take in 3 arguments and returns the maximum of the three numbers


```python
biggest_change = ...
biggest_change
```

```python
grader.check("q2_2_1")
```

<!-- END QUESTION -->

**Question 2.** Which of the three majors had the **smallest** absolute difference? Assign `smallest_change_major` to 1, 2, or 3 where each number corresponds to the following major:

1: Gender and Women's Studies  
2: Linguistics  
3: Rhetoric

Choose the number that corresponds to the major with the smallest absolute difference.

You should be able to answer by rough mental arithmetic, without having to calculate the exact value for each major. 


```python
smallest_change_major = ...
smallest_change_major
```

```python
grader.check("q2_2_2")
```

**Question 3.**  For each major, define the "relative change" to be the following: $\large{\frac{\text{absolute difference}}{\text{value in 2008-2009}} * 100}$ 

Fill in the code below such that `gws_relative_change`, `linguistics_relative_change` and `rhetoric_relative_change` are assigned to the relative changes for their respective majors.


```python
gws_relative_change = (abs(...) / 17) * 100
linguistics_relative_change = ...
rhetoric_relative_change = ...
gws_relative_change, linguistics_relative_change, rhetoric_relative_change
```

```python
grader.check("q2_2_3")
```

**Question 4.** Assign `biggest_rel_change_major` to 1, 2, or 3 where each number corresponds to to the following: 

1: Gender and Women's Studies  
2: Linguistics  
3: Rhetoric

Choose the number that corresponds to the major with the biggest relative change.


```python
# Assign biggest_rel_change_major to the number corresponding to the major with the biggest relative change.
biggest_rel_change_major = ...
biggest_rel_change_major
```

```python
grader.check("q2_2_4")
```

## Acknowledgements

<!-- BEGIN QUESTION -->

**Acknowledgements:**
If you received help from any person other than your instructor, then describe here what help you received and from whom. 
If not, then say 'None'.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

It is forbidden to: 
* Use the work or programs from students who took this class (in any of its versions at any university).
* Use websites that claim to have solutions for this class (in any of its versions at any university).
* Directly copy from a classmate.

**Did you do anything forbidden while working on this assignment?**

Assign `did_forbidden` to either `True` or `False`.

```python
did_forbidden = ...
```

```python
grader.check("q_forbidden")
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

After exporting, submit the .zip file in Canvas

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for q2_1_1 (points: None)
```
>>> 1 <= names_q1 <= 4
True
```


Tests for q2_1_2 (points: None)
```
>>> 1 <= names_q2 <= 4
True
```


Tests for q2_1_3 (points: None)
```
>>> 1 <= names_q3 <= 4
True
```


Tests for q2_2_1 (points: None)
```
>>> isinstance(biggest_change, (int, float))
True
```


Tests for q2_2_2 (points: None)
```
>>> isinstance(smallest_change_major, (int, float))
True
```


Tests for q2_2_3 (points: None)
```
>>> isinstance(gws_relative_change, (int, float))
True
```
```
>>> isinstance(linguistics_relative_change, (int, float))
True
```
```
>>> isinstance(rhetoric_relative_change, (int, float))
True
```


Tests for q2_2_4 (points: None)
```
>>> isinstance(biggest_rel_change_major, (int, float))
True
```


Tests for q_forbidden (points: None)
```
>>> did_forbidden in [True, False]
True
```
